---
title_meta: Capitolo 1
title: Introduzione alle basi
description: >-
  Muovi i primi passi con R. In questo capitolo imparerai a usare la console
  come una calcolatrice e ad assegnare valori alle variabili. Inoltre,
  conoscerai anche i principali tipi di dati in R. Iniziamo!
free_preview: true
---

## Come funziona

```yaml
type: NormalExercise
key: 15d729634a
xp: 100
skills:
  - 1
```

Nell’editor devi digitare il codice per risolvere gli esercizi. Quando lo invii, ogni riga viene interpretata ed eseguita da R, e riceverai un messaggio che ti dirà se è corretta o meno. L’output viene visualizzato nella console.

In R, il simbolo `#` serve per aggiungere commenti, in modo da permettere a te e agli altri di capire di cosa tratta il codice. Un po’ come su Twitter! I commenti non vengono eseguiti come codice R, quindi non influenzeranno il risultato. Ad esempio, <span translate="no">_Calculate 3 + 4_</span> nell'editor è solo un commento.

Puoi anche eseguire i comandi R direttamente nella console. È un buon modo per sperimentare con il codice, in quanto i comandi digitati nella console non saranno valutati come parte dell'esercizio.

`@instructions`
- Nell'editor sono già presenti alcuni esempi di codice. Riesci a distinguere quali sono le righe di codice R e quali sono i commenti?
- Aggiungi una riga di codice che calcola la somma di 6 e 12 e invialo.

`@hint`
Basta aggiungere una riga di codice R che calcola la somma di 6 e 12, proprio come nell'esempio!

`@pre_exercise_code`
```{r}
# no pec

```

`@sample_code`
```{r}
# Calculate 3 + 4
3 + 4

# Calculate 6 + 12

```

`@solution`
```{r}
# Calculate 3 + 4
3 + 4

# Calculate 6 + 12
6 + 12
```

`@sct`
```{r}
ex() %>% check_output_expr("18", missing_msg = "Assicurati di aggiungere `6 + 12` su una nuova riga. Non iniziare la riga con un `#`, altrimenti il tuo codice R non viene eseguito!")

success_msg("Fantastico! Vedi come la console mostra il risultato del codice R che hai inviato? Ora che hai familiarità con l'interfaccia, passiamo al lavoro con R!")
```

---

## Aritmetica con R

```yaml
type: NormalExercise
key: 720745eda5
xp: 100
skills:
  - 1
```

Nella sua forma più elementare, R può essere utilizzato come una semplice calcolatrice. Considera i seguenti operatori aritmetici:

- Somma: `+`
- Sottrazione: `-`
- Moltiplicazione: `*`
- Divisione: `/`
- Potenza: `^`
- Funzione modulo: `%%`

Gli ultimi due potrebbero richiedere qualche spiegazione:

- L'operatore `^` eleva il numero alla sua sinistra alla potenza del numero alla sua destra: ad esempio `3^2` è 9.
- L’operatore modulo restituisce il resto della divisione del numero a sinistra per il numero alla sua destra, ad esempio 5 modulo 3 o `5 %% 3` è 2.

Con queste conoscenze, segui le istruzioni per completare l'esercizio.

`@instructions`
- Digita `2^5` nell'editor per calcolare 2 alla potenza di 5.
- Digita `28 %% 6` per calcolare 28 modulo 6.
- Invia la risposta e osserva l'output di R nella console.
- Ricorda che il simbolo `#` è utilizzato per aggiungere commenti al codice R.

`@hint`
Un altro esempio dell’operatore modulo: `9 %% 2` è uguale a `1`.

`@pre_exercise_code`
```{r}
# no pec
```

`@sample_code`
```{r}
# An addition
5 + 5 

# A subtraction
5 - 5 

# A multiplication
3 * 5

 # A division
(5 + 5) / 2 

# Exponentiation


# Modulo

```

`@solution`
```{r}
# An addition
5 + 5

# A subtraction
5 - 5 

# A multiplication
3 * 5

 # A division
(5 + 5) / 2 

# Exponentiation
2 ^ 5

# Modulo
28 %% 6
```

`@sct`
```{r}
msg = "Non rimuovere gli altri esempi aritmetici!"
ex() %>% check_output_expr("2^5", missing_msg = "L'esempio di esponenziazione non è corretto. Scrivi `2 ^ 5` su una nuova riga.")

ex() %>% check_output_expr("28 %% 6", missing_msg = "Sembra esserci un problema con l'esempio del modulo. Scrivi `28 %% 6` su una nuova riga.")
success_msg("Ottimo! Passa all'esercizio successivo.")
```

---

## Assegnazione delle variabili

```yaml
type: NormalExercise
key: 5f200ffd43
xp: 100
skills:
  - 1
```

Un concetto fondamentale della programmazione (statistica) è chiamato **variabile**. 

Una variabile ti permette di memorizzare un valore (ad es. 4) o un oggetto (ad es. la descrizione di una funzione) in R. In seguito potrai utilizzare il nome di questa variabile per accedere facilmente al valore o all'oggetto memorizzato al suo interno. 

Puoi assegnare il valore 4 alla variabile `my_var` con il comando

```
my_var <- 4
```

`@instructions`
Ora tocca a te: completa il codice nell'editor in modo da assegnare il valore 42 alla variabile `x`, poi invia la risposta. Nota che, quando chiedi a R di stampare `x`, viene visualizzato il valore 42.

`@hint`
Guarda come il valore 4 è stato assegnato a `my_var` nello svolgimento dell’esercizio. Fai la stessa cosa nell'editor, ma ora assegna 42 alla variabile `x`.

`@pre_exercise_code`
```{r}
# no pec
```

`@sample_code`
```{r}
# Assign the value 42 to x
x <- 

# Print out the value of the variable x
x
```

`@solution`
```{r}
# Assign the value 42 to x
x <- 42

# Print out the value of the variable x
x
```

`@sct`
```{r}
ex() %>% check_object("x", undefined_msg = "Assicurati di definire una variabile `x`.")  %>% check_equal(incorrect_msg = "Assicurati di assegnare il valore corretto a `x`.")

success_msg("Ottimo lavoro! Hai notato che R non stampa il valore di una variabile nella console quando fai l'assegnazione? `x <- 42` non ha generato alcun output, perché R presume che avrai bisogno di questa variabile in futuro. Altrimenti non avresti memorizzato il valore in una variabile, giusto? Procedi al prossimo esercizio!")
```

---

## Assegnazione delle variabili (2)

```yaml
type: NormalExercise
key: c5944b90eb
xp: 100
skills:
  - 1
```

Supponiamo di avere un cesto di frutta con cinque mele. Come Data Analyst in formazione, vuoi memorizzare il numero di mele in una variabile con il nome `my_apples`.

`@instructions`
- Digita il seguente codice nell'editor: `my_apples <- 5`. In questo modo assegnerai il valore 5 alla variabile `my_apples`.
- Scrivi `my_apples` sotto il secondo commento. In questo modo verrà stampato il valore di `my_apples`.
- Invia la risposta e osserva l'output: verrà stampato il numero 5. Ora R avrà associato il valore 5 alla variabile `my_apples`.

`@hint`
Ricorda che se vuoi assegnare un numero o un oggetto a una variabile in R, puoi usare l'operatore di assegnazione `<-`. In alternativa, puoi utilizzare `=`, ma `<-` è ampiamente preferito dalla comunità di R.

`@pre_exercise_code`
```{r}
# no pec
```

`@sample_code`
```{r}
# Assign the value 5 to the variable my_apples


# Print out the value of the variable my_apples

```

`@solution`
```{r}
# Assign the value 5 to the variable my_apples
my_apples <- 5

# Print out the value of the variable my_apples
my_apples
```

`@sct`
```{r}
ex() %>% check_object("my_apples", undefined_msg = "Per favore, assicurati di definire una variabile `my_apples`.") %>% check_equal(incorrect_msg = "Assicurati di assegnare il valore corretto a `my_apples`.")

ex() %>% check_output_expr("my_apples", missing_msg = "Hai detto esplicitamente a R di stampare la variabile `my_apples` nella console?")

success_msg("Ottimo! Continua con l'esercizio successivo!")
```

---

## Assegnazione delle variabili (3)

```yaml
type: NormalExercise
key: 1c1bd25045
xp: 100
skills:
  - 1
```

Ogni cesto di frutta ha bisogno di arance, quindi decidi di aggiungere sei arance. In qualità di Data Analyst, il tuo istinto è creare immediatamente la variabile `my_oranges` e assegnarle il valore 6. Successivamente, dovrai calcolare quanti pezzi di frutta hai in totale. Ora che hai assegnato un nome a questi valori, puoi scrivere il codice in modo più comprensibile: 

```
my_apples + my_oranges
```

`@instructions`
- Assegna a `my_oranges` il valore di 6.
- Aggiungi le variabili `my_apples` e `my_oranges`, e fai in modo che R stampi il risultato.
- Assegna il risultato della somma di `my_apples` e `my_oranges` a una nuova variabile `my_fruit`.

`@hint`
`my_fruit` è semplicemente la somma di `my_apples` e `my_oranges`. Puoi usare l'operatore `+` per sommare i due valori e `<-` per assegnare il valore alla variabile `my_fruit`.

`@pre_exercise_code`
```{r}
# no pec
```

`@sample_code`
```{r}
# Assign a value to the variables my_apples and my_oranges
my_apples <- 5


# Add these two variables together


# Create the variable my_fruit

```

`@solution`
```{r}
# Assign a value to the variables my_apples and my_oranges
my_apples  <- 5
my_oranges <- 6

# Add these two variables together
my_apples + my_oranges

# Create the variable my_fruit
my_fruit <- my_apples + my_oranges
```

`@sct`
```{r}
msg <- "Hai usato `my_fruit <- my_apples + my_oranges` per creare la variabile `my_fruit`?"
ex() %>% check_object("my_apples") %>% check_equal(incorrect_msg = "Mantieni la riga che assegna 5 a `my_apples`.")

ex() %>% check_object("my_oranges") %>% check_equal(incorrect_msg = "Mantieni la riga che assegna 6 a `my_oranges`.")

ex() %>% check_output_expr("my_apples + my_oranges",missing_msg = "Assicurati di stampare il risultato dell\'aggiunta di `my_apples` e `my_oranges`. L\'esempio di codice nella descrizione ti dà già la risposta a questa istruzione!")

ex() %>% check_object("my_fruit", undefined_msg = msg)  %>% check_equal(incorrect_msg = msg)
success_msg("Ben fatto! Il grande vantaggio di fare calcoli con le variabili è la riusabilità. Se cambi solo `my_apples` per essere uguale a 12 invece di 5 e riesegui lo script, `my_fruit` si aggiornerà automaticamente. Continua al prossimo esercizio.")
```

---

## Mele e arance

```yaml
type: NormalExercise
key: 915fcc7c99
xp: 100
skills:
  - 1
```

La saggezza popolare dice di non sommare mele e arance. Ma è proprio quello che hai fatto poco fa! :-) Le variabili `my_apples` e `my_oranges` contenevano entrambe un valore numerico nell'esercizio precedente. In R l'operatore `+` funziona con variabili numeriche. Se invece provassi a sommare <span translate="no">"apples"</span> e <span translate="no">"oranges"</span> assegnando un valore di testo a `my_oranges` (vedi nell'editor), staresti provando a sommare un valore numerico e uno di tipo carattere alla variabile `my_fruit`. Questo non è possibile in R.

`@instructions`
- Invia la risposta e leggi il messaggio di errore. Assicurati di capire perché non ha funzionato.
- Sistema il codice in modo che R sappia che hai 6 arance e quindi un cesto di frutta con 11 pezzi di frutta.

`@hint`
Devi assegnare alla variabile `my_oranges` il valore numerico `6`, non il valore di tipo carattere `"six"`. Nota che le virgolette sono utilizzate per indicare che `"six"` è un carattere (testo).

`@pre_exercise_code`
```{r}
# no pec
```

`@sample_code`
```{r}
# Assign a value to the variable my_apples
my_apples <- 5 

# Fix the assignment of my_oranges
my_oranges <- "six" 

# Create the variable my_fruit and print it out
my_fruit <- my_apples + my_oranges 
my_fruit
```

`@solution`
```{r}
# Assign a value to the variable my_apples
my_apples <- 5  

# Fix the assignment of my_oranges
my_oranges <- 6

# Create the variable my_fruit and print it out
my_fruit <- my_apples + my_oranges 
my_fruit
```

`@sct`
```{r}
ex() %>% check_error(incorrect_msg = "Puoi farlo impostando la variabile `my_oranges` su un valore numerico, non una stringa!")

ex() %>% check_object("my_apples") %>% check_equal(incorrect_msg = "Assicurati che `my_apples` contenga ancora `5`.")

ex() %>% check_object("my_oranges") %>% check_equal(incorrect_msg = "Assicurati che `my_oranges` sia uguale a `6`.")

ex() %>% check_object("my_fruit") %>% check_equal(incorrect_msg = "Il valore di `my_fruit` non è corretto. Dovrebbe essere 11, la somma di `my_apples` e `my_oranges`.")

ex() %>% check_object("my_fruit")%>% check_or(
  check_code(.,"my_fruit\\s*<-\\s*my_apples\\s*\\+\\s*my_oranges", missing_msg = "Hai creato `my_fruit` come la somma delle variabili `my_apples` e `my_oranges`?", append=F),
  check_code(.,"my_fruit\\s*<-\\s*my_oranges\\s*\\+\\s*my_apples", missing_msg = "Hai creato `my_fruit` come la somma delle variabili `my_apples` e `my_oranges`?", append=F)
 )
ex() %>% check_output_expr("my_fruit", missing_msg = "Non rimuovere la riga che stampa `my_fruit`.")

success_msg("Fantastico, continua così! Passa al prossimo esercizio.")
```

---

## Principali tipi di dati in R

```yaml
type: NormalExercise
key: 0f23107394
xp: 100
skills:
  - 1
```

R lavora con numerosi tipi di dati. Alcuni dei tipi più semplici per iniziare sono:

- I valori decimali come `4.5` sono chiamati **numerici**.
- I numeri interi come `4` sono chiamati **numeri interi**. Gli interi sono anche numerici.
- I valori booleani (`TRUE` o `FALSE`) sono chiamati **logici**.
- I valori di testo (o stringhe) sono chiamati **caratteri**.

Nota che le virgolette nell'editor indicano che `"some text"` è una stringa.

`@instructions`
Cambia il valore di:

- `my_numeric` a `42`.
- `my_character` a `"universe"`. Le virgolette indicano che `"universe"` è un carattere.
- `my_logical` a `FALSE`.

Nota che R distingue anche tra maiuscole e minuscole!

`@hint`
Sostituisci i valori nell'editor con quelli forniti nell'esercizio. Ad esempio:
`my_numeric <- 42` assegna il valore 42 alla variabile `my_numeric`.

`@pre_exercise_code`
```{r}
# no pec
```

`@sample_code`
```{r}
# Change my_numeric to be 42
my_numeric <- 42.5

# Change my_character to be "universe"
my_character <- "some text"

# Change my_logical to be FALSE
my_logical <- TRUE
```

`@solution`
```{r}
# Change my_numeric to be 42
my_numeric <- 42

# Change my_character to be "universe"
my_character <- "universe"

# Change my_logical to be FALSE
my_logical <- FALSE
```

`@sct`
```{r}
ex() %>% check_object("my_numeric") %>% check_equal(incorrect_msg = "Hai correttamente cambiato la dichiarazione di `my_numeric` in modo che contenga il valore 42?")

ex() %>% check_object("my_character") %>% check_equal(incorrect_msg = "Hai correttamente cambiato `my_character` in `\"universe\"`? Non dimenticare le virgolette!")

ex() %>% check_object("my_logical") %>% check_equal(incorrect_msg = "Hai correttamente cambiato `my_logical` in `FALSE`? Tutte le lettere di `FALSE` devono essere maiuscole!")

success_msg("Ottimo lavoro! Continua con l'esercizio successivo.")
```

---

## Qual è il tipo di dati?

```yaml
type: NormalExercise
key: 99b549229d
xp: 100
skills:
  - 1
```

Ricordi che quando hai sommato `5 + "six"`, si è verificato un errore dovuto a una mancata corrispondenza tra i tipi di dati? Puoi evitare queste situazioni controllando prima il tipo di dati di una variabile. Puoi farlo con la funzione `class()`, come mostra il codice nell'editor.

`@instructions`
Completa il codice nell'editor e stampa anche `my_character` e `my_logical`.

`@hint`
Il codice per stampare il tipo di dati di `my_numeric` è già presente; fai lo stesso per `my_character` e `my_logical`.

`@pre_exercise_code`
```{r}
# no pec
```

`@sample_code`
```{r}
# Declare variables of different types
my_numeric <- 42
my_character <- "universe"
my_logical <- FALSE 

# Check class of my_numeric
class(my_numeric)

# Check class of my_character


# Check class of my_logical

```

`@solution`
```{r}
# Declare variables of different types:
my_numeric <- 42
my_character <- "universe"
my_logical <- FALSE

# Check class of my_numeric
class(my_numeric)

# Check class of my_character
class(my_character)

# Check class of my_logical
class(my_logical)
```

`@sct`
```{r}
msg <- "Non cambiare la dichiarazione delle variabili!"
ex()  %>% check_object("my_numeric", undefined_msg = msg)  %>%  check_equal(incorrect_msg = msg)

ex()  %>% check_object("my_character", undefined_msg = msg)  %>%  check_equal(incorrect_msg = msg)

ex()  %>% check_object("my_logical", undefined_msg = msg)  %>%  check_equal(incorrect_msg = msg)

patt <- "Hai incluso `class(%1$s)` per stampare il tipo di dato di `%1$s`?"
ex() %>% check_output_expr("class(my_numeric)",missing_msg = "Non rimuovere il codice che stampa il tipo di `my_numeric`.")

ex() %>% check_output_expr("class(my_character)",missing_msg = sprintf(patt, "my_character"))

ex() %>% check_output_expr("class(my_logical)",missing_msg = sprintf(patt, "my_logical"))
success_msg("Congratulazioni! Questo era l'ultimo esercizio di questo capitolo. Vai al prossimo capitolo per immergerti nel mondo dei vettori!")
```
