--- 
title_meta  : Chapter 6
title       : Lists
description : Lists, as opposed to vectors, can hold components of different types, just like your to-do list at home or at work. This intro to R chapter will teach you how to create, name and subset these lists.




--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:aafcdee5a9
## You know lots of things already

Congratulations! At this point in the course you are already familiar with vectors, factors, matrices, and data frames. Pretty sweet for an R newbie, right? ;-)

Before you start learning about lists, let's make sure you've understood these variable types. Which one of the following statements is false?

*** =instructions

* Vectors are one dimensional variables that can hold numeric, character or logical values.
* The elements in a vector or matrix must all have the same data type.
* Factors are one dimensional variables that hold categorical data.
* Within a row of a data frame, each element must have the same data type.
* Matrices and data frames are both two dimensional variables.

*** =hint

What are the restrictions on the data that can be stored in data frames?

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
not_wrong_msg <- "Sorry, that statement is true."
correct_msg <- "Very insightful! A data frame can have different data types within a row, but must have the same data type within a column."
test_mc(
  4,
  feedback_msgs = c(rep.int(not_wrong_msg, 3), correct_msg, not_wrong_msg)
)
```




--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:09c072ea7a
## What are lists?

Between numeric, character and logical variables, matrices, factors, and data frames, that is six variable types that you understand. In total, R has about twenty variable types, so you know about a good portion of the language already.

Restricting R to only twenty data types makes the language easier to work with. However, sometimes you need more flexibility, and this is where lists become useful.

Lists are a container variable.  That is, they are a special type of vector that can store other variable types.  Each element of a list can be a vector, a data frame, or any other variable type. Lists can even contain other lists!

The twenty data types are like Lego bricks. In the same way that Lego bricks can be combined into an infinite number of shapes, by combining R variables together using lists, you can make an infinite number of different data structures.

Which one of the following statements is true?

*** =instructions

* Lists let you combine other variable types into one variable.
* Lists can contain any other data type except other lists.
* R has hundreds of data types to provide flexibility.
* Lists are a type of matrix.

*** =hint

* Lists can contain lists.
* R has around 20 data types.
* Lists are vectors, but not matrices.

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
msgs <- list(
  correct = "Marvellous! Lists are used to combine other variables into a single variable.",
  contains = "Lists can contain other lists.",
  hundreds = "R only has around 20 data types.",
  matrix = "Lists are a type of vector, but not a type of matrix."
)

test_mc(
  1,
  feedback_msgs = with(msgs, c(correct, contains, hundreds, matrix))
)
```



--- type:NormalExercise xp:100 skills:1 key:4beee9cb532c889903218b49b83ab5ef133eac83
## Creating a list

Time to create your first list! To construct a list you use the function [`list()`](http://www.rdocumentation.org/packages/base/functions/list). You use it in the same way that you used [`c()`](http://www.rdocumentation.org/packages/base/functions/c) to create `numeric`, `logical`, and `character` vectors.

```
my_list <- list(comp1, comp2 ...)
```

The arguments to the list function are the list components. Remember, these components can be any variable types that you like, including matrices, vectors, or other lists.

*** =instructions

The variables `my_vector`, `my_matrix` and `my_df` are shown in your script.

* Construct a list, named `my_list`, that contains `my_vector`, `my_matrix` and `my_df` as list components, in that order.

*** =hint
Use the [`list()`](http://www.rdocumentation.org/packages/base/functions/list) function with `my_vector`, `my_matrix` and `my_df` as arguments separated by a comma.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vector with numerics from 1 up to 10
my_vector <- 1:10 

# Matrix with numerics from 1 up to 9
my_matrix <- matrix(1:9, ncol = 3)

# First 10 elements of the built-in data frame mtcars
my_df <- mtcars[1:10, ]

# Construct list with these different elements:
my_list <- 
```

*** =solution
```{r}
# Vector with numerics from 1 up to 10
my_vector <- 1:10 

# Matrix with numerics from 1 up to 9
my_matrix <- matrix(1:9, ncol = 3)

# First 10 elements of the built-in data frame mtcars
my_df <- mtcars[1:10, ]

# Construct list with these different elements:
my_list <- list(my_vector, my_matrix, my_df)
```

*** =sct
```{r}
msg = "Do not remove or change the definition of the variables `my_vector`, `my_matrix` or `my_df`!"
test_object("my_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("my_matrix", undefined_msg = msg, incorrect_msg = msg)
test_object("my_df", undefined_msg = msg, incorrect_msg = msg)
test_object("my_list",
            incorrect_msg = "It looks like `my_list` does not contain the correct elements. Make sure to pass the variables `my_vector`, `my_matrix` and `my_df` to the `list()` function, separated by commas, in this order.")
success_msg("Wonderful! Each element of the list is passed to `list()` separated by commas.")
```


--- type:NormalExercise xp:100 skills:1 key:89dd0126568b1ff5a84033c571907a8a282245e4
## Creating a named list

Well done, you're on a roll!

You want to avoid not knowing or remembering what the components of your list stand for. That is why you should give names to them: 

```
my_list <- list(name1 = your_comp1, 
                name2 = your_comp2)
``` 

This creates a list with components that are named `name1`, `name2`, and so on. If you want to name your lists after you've created them, you can use the `names()` function as you did with vectors. The following commands are fully equivalent to the assignment above:

```
my_list <- list(your_comp1, your_comp2)
names(my_list) <- c("name1", "name2")
```

*** =instructions
- Change the code of the previous exercise (see editor) by adding names to the components. Use for `my_vector` the name `vec`, for `my_matrix` the name `mat` and for `my_df` the name `df`. 
- Print out `my_list` so you can inspect the output.

*** =hint
The first method of assigning names to your list components is the easiest. It starts like this: 
```
my_list <- list(vec = my_vector)
```
Add the other two components in a similar fashion.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vector with numerics from 1 up to 10
my_vector <- 1:10 

# Matrix with numerics from 1 up to 9
my_matrix <- matrix(1:9, ncol = 3)

# First 10 elements of the built-in data frame mtcars
my_df <- mtcars[1:10, ]

# Adapt list() call to give the components names
my_list <- list(my_vector, my_matrix, my_df)

# Print out my_list

```

*** =solution
```{r}
# Vector with numerics from 1 up to 10
my_vector <- 1:10 

# Matrix with numerics from 1 up to 9
my_matrix <- matrix(1:9, ncol = 3)

# First 10 elements of the built-in data frame mtcars
my_df <- mtcars[1:10, ]

# Adapt list() call to give the components names
my_list <- list(vec = my_vector, mat = my_matrix, df = my_df)

# Print out my_list
my_list
```

*** =sct
```{r}
msg = "Do not remove or change the definiton of the variables `my_vector`, `my_matrix` or `my_df`!"
test_object("my_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("my_matrix", undefined_msg = msg, incorrect_msg = msg)
test_object("my_df", undefined_msg = msg, incorrect_msg = msg)
test_object("my_list",
            incorrect_msg = "It looks like `my_list` does not contain the correct elements.")
test_object("my_list", eq_condition = "equal",
            incorrect_msg = "It looks like `my_list` does not contain the correct naming for the components. Make sure you use the names `vec`, `mat` and `df`, respectively. Check out the hint if you're struggling.");
test_output_contains("my_list", incorrect_msg = "Don't forget to print `my_list` to the console! Simply add `my_list` on a new line in the editor.")
success_msg("Great! Not only do you know how to construct lists now, you can also name them; a skill that will prove most useful in practice. Continue to the next exercise.")
```


--- type:NormalExercise xp:100 skills:1 key:19fd17cc792ef870c1558b3a9bae08c1d1e3acae
## Creating a named list (2)

Being a huge movie fan (remember your job at Disney), you decide to start storing information on good movies with the help of lists. 

Start by creating a list for the movie "The Shining".

*** =instructions

The variables `mov`, `act` and `rev` have been pre-defined in your R workspace.

* Complete the code on the right to create `shining_list`; it should contain three elements:
    - moviename: a character string with the movie title (stored in `mov`)
    - actors: a vector with the main actors' names (stored in `act`)
    - reviews: a data frame that contains some reviews (stored in `rev`)

Do not forget to name the list components accordingly (names are moviename, actors and reviews).

*** =hint
`shining_list <- list(moviename = mov)` is only part of the solution; it's your job to also add `act` and `rev` to the list, with the appropriate names.

*** =pre_exercise_code
```{r}
mov <- "The Shining"
act <- c("Jack Nicholson", "Shelley Duvall", "Danny Lloyd", "Scatman Crothers", "Barry Nelson")
rev <- data.frame(
  scores = c(4.5,4,5), 
  sources = c("IMDb1","IMDb2","IMDb3"), 
  comments = c(
    "Best Horror Film I Have Ever Seen", 
    "A truly brilliant and scary film from Stanley Kubrick", 
    "A masterpiece of psychological horror"
  )
)
rm(scores, sources, comments)
```

*** =sample_code
```{r}
# The variables mov, act and rev are available
mov
act
rev

# Finish the code to build shining_list
shining_list <- list(moviename = mov, ___ = ___, ___ = ___)
```

*** =solution
```{r}
# The variables mov, act and rev are available
mov
act
rev

# Finish the code to build shining_list
shining_list <- list(moviename = mov, actors = act, reviews = rev)
```

*** =sct
```{r}
msg = "Do not remove or change the definition of the pre-set variables!"
lapply(c("mov", "rev", "act"), test_object, undefined_msg = msg, incorrect_msg = msg)
test_object("shining_list",
            incorrect_msg = "It looks like `shining_list` does not contain the correct elements: the first element should be `mov`, the second element `act`, and the third `rev`.")
test_object("shining_list", eq_condition = "equal",
            incorrect_msg = "It looks like `shining_list` does not contain the correct naming for the components: name the first element `moviename`, the second element `actors`, and the third element `reviews`.");
success_msg("Wonderful! You now know how to construct and name lists. As in the previous chapters, let's look at how to select elements for lists. Head over to the next exercise")
```


--- type:NormalExercise xp:100 skills:1 key:1ef3278944562caef64b9927dd2ddb6ee52334cd
## Selecting elements from a list

You've seen how to use square brackets to select elements from vectors, matrices, and data frames. Now take a look at how you can use them to select elements from lists.

Consider this "list": a jar containing three "vectors" of almonds. 

![a jar containing three bags of almonds](http://s3.amazonaws.com/assets.datacamp.com/production/course_2977/datasets/intro-to-r-list-indexing-complete-list.jpg)

If you use square brackets to select the second element, you get another jar, containing one vector of almonds.  That is, using square bracket indexing will always return another list.

```{r}
jar[2]
```

![a jar containing one bag of almonds](http://s3.amazonaws.com/assets.datacamp.com/production/course_2977/datasets/intro-to-r-list-indexing-single-square-bracket.jpg)

Just as with other vectors, you can use positive numbers to select elements to keep, negative numbers to select elements to remove, or a logical vector to keep elements where that vector is `TRUE`.  If the elements of your list are named, you can also select elements by those names.

*** =instructions

The variable `shining_list`, containing data about the movie "The Shining", has been pre-defined for you. In each of the following exercises, use square bracket indexing, `[`.

- Use positive number indexing to select the second and third elements from `shining_list`.
- Use negative number indexing to select everything but the third element of `shining_list`.
- Use character indexing to select the `"moviename"` element of `shining_list`.

*** =hint
- To select the elements using positive indexing, use code like `[c(1, 3, 5)]`.
- To select the elements using negative indexing, use code like `[c(-1, -3, -5)]`.
- To select the elements using character indexing, use code like `[c("alpha", "charlie", "echo")]`.

*** =pre_exercise_code
```{r}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/shining_list.RData"))
```

*** =sample_code
```{r}
# shining_list is already pre-loaded in the workspace
shining_list

# Select 2nd and 3rd element using positive indexing
shining_list[]

# Select all but 3rd element using negative indexing


# Select moviename element using character indexing

```

*** =solution
```{r}
# shining_list is already pre-loaded in the workspace
shining_list

# Select 2nd and 3rd element using positive indexing
shining_list[2:3]

# Select all but 3rd element using negative indexing
shining_list[-3]

# Select moviename element using character indexing
shining_list["moviename"]
```

*** =sct
```{r}
predef_msg <- "Do not remove or change the definition of `shining_list`, which is pre-loaded in the workspace!"
ex() %>% {
  check_object(., "shining_list", undefined_msg = predef_msg) %>%
    check_equal(incorrect_msg = predef_msg)
  check_output_expr(., "shining_list[2:3]")
  check_output_expr(., "shining_list[-3]")
  check_output_expr(., 'shining_list["moviename"]')
}
success_msg("Great! Selecting elements from lists with square brackets always returns another list.")
```



--- type:NormalExercise lang:r xp:100 skills:1 key:d355af51a9
## Selecting elements from a list (2)

Sometimes, rather than wanting another list, you want to select the contents of one element of your list. For example, rather than wanting a jar containing a bag of almonds, you just want a bag of almonds.  To do this, you use double square brackets, `[[ ]]`.

```{r}
jar[[2]]
```

![a bag of almonds](http://s3.amazonaws.com/assets.datacamp.com/production/course_2977/datasets/intro-to-r-list-indexing-double-square-bracket.jpg)

For double square bracket indexing, you can only select a single element from the list, so the value in between the brackets must be a single positive number, or a single string naming an element.

*** =instructions

The variable `shining_list` has been pre-defined for you. In each of the following exercises, use double square bracket indexing, `[[`.

- Use positive number indexing to select the third element from `shining_list`.
- Use character indexing to select the `"actors"` element of `shining_list`.


*** =hint

* To select the contents of the fifth element of a list using positive number indexing, you would use `[[5]]`.
* To select the contents of the `"bravo"` element of a list using character indexing, you would use `[["bravo"]]`.

*** =pre_exercise_code
```{r}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/shining_list.RData"))
```

*** =sample_code
```{r}
# shining_list is already pre-loaded in the workspace
shining_list

# Select 3rd element using positive indexing
shining_list[[]]

# Select actors element using character indexing

```

*** =solution
```{r}
# shining_list is already pre-loaded in the workspace
shining_list

# Select 3rd element using positive indexing
shining_list[[3]]

# Select actors element using character indexing
shining_list[["actors"]]
```

*** =sct
```{r}
predef_msg <- "Do not remove or change the definition of `shining_list`, which is pre-loaded in the workspace!"
ex() %>% {
  check_object(., "shining_list", undefined_msg = predef_msg) %>%
    check_equal(incorrect_msg = predef_msg)
  check_output_expr(., "shining_list[[3]]")
  check_output_expr(., 'shining_list[["actors"]]')
}
success_msg("Magnificent! Selecting an element from a list with double square brackets returns the contents of that list element.")
```


--- type:NormalExercise lang:r xp:100 skills:1 key:ec42b8889a
## Selecting elements from a list (3)

All these exercises selecting elements are making you hungry.  This time you want to select some almonds!  To do this, you need to chain your selections together.  This is easier than it might sound.  All you have to do is to write the square brackets one after the other.  For example, to select the second "vector" of almonds, you used

```{r}
jar[[2]]
```

Now to select a portion of the first five almonds from that "vector", you could use the following code.

```{r}
jar[[2]][1:5]
```

![five almonds](http://s3.amazonaws.com/assets.datacamp.com/production/course_2977/datasets/intro-to-r-list-indexing-chained-indexing.jpg)

To recap what happened, you used double square brackets to access the vector stored in the list, then you followed it with single square brackets to access the elements in the vector.

*** =instructions

The variable `shining_list`, containing data about the movie "The Shining", has been pre-defined for you.

* Select the first, third, and fifth actors from `shining_list`.
* Select all of the first row of the reviews from `shining_list`. 

*** =hint

* To select the contents of the first element of a list, you would use `[[1]]`.
* To select the second and fourth elements of a vector, you would use `[c(2, 4)]`.
* To select the fifth row of a data frame you would use `[5, ]`.

*** =pre_exercise_code
```{r}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/shining_list.RData"))
```

*** =sample_code
```{r}
# shining_list is already pre-loaded in the workspace
shining_list

# Select the 1st, 3rd, and 5th actors
shining_list[[]][]

# Select the first review

```

*** =solution
```{r}
# shining_list is already pre-loaded in the workspace
shining_list

# Select the 1st, 3rd, and 5th actors
shining_list[["actors"]][c(1, 3, 5)]

# Select the first review
shining_list[["reviews"]][1, ]
```

*** =sct
```{r}
predef_msg <- "Do not remove or change the definition of `shining_list`, which is pre-loaded in the workspace!"
ex() %>% {
  check_object(., "shining_list", undefined_msg = predef_msg) %>%
    check_equal(incorrect_msg = predef_msg)
  check_output_expr(., 'shining_list[["actors"]][c(1, 3, 5)]')
  check_output_expr(., 'shining_list[["reviews"]][1, ]')
}
success_msg("You superstar! Square brackets can be chained together to select elements inside other elements.")
```
--- type:NormalExercise xp:100 skills:1 key:ca3f7b71504ff93940cf65889d406a58c5b0019c
## Adding more movie information to the list

Being proud of your first list, you shared it with the members of your movie hobby club. However, one of the senior members, a guy named M. McDowell, noted that you forgot to add the release year. Given your ambitions to become next year's president of the club, you decide to add this information to the list. 

To conveniently add elements to lists you can use the [`c()`](http://www.rdocumentation.org/packages/base/functions/c) function, that you also used to build vectors:

```
ext_list <- c(my_list , my_val)
``` 

This will simply extend the original list, `my_list`, with the component `my_val`. This component gets appended to the end of the list.
If you want to give the new list item a name, you just add the name as you did before: 

```
ext_list <- c(my_list, my_name = my_val)
```

*** =instructions
- Complete the code below such that an item named `year` is added to the `shining_list` with the value 1980. Assign the result to `shining_list_full`.
- Finally, have a look at the structure of `shining_list_full` with the [`str()`](http://www.rdocumentation.org/packages/utils/functions/str) function.

*** =hint
Have a look at the example code in the exercise assignment. Maybe this can help you start:
```
shining_list <- c(shining_list, ...)
```
You still have to add some code where the three dots are.

*** =pre_exercise_code
```{r, eval = FALSE}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/shining_list.RData"))
```

*** =sample_code

```{r}
# shining_list, the list containing movie name, actors and reviews, is pre-loaded in the workspace

# We forgot something; add the year to shining_list
shining_list_full <- 

# Have a look at shining_list_full

```

*** =solution

```{r}
# shining_list, the list containing movie name, actors and reviews, is pre-loaded in the workspace

# Use c() to add a year to shining_list
shining_list_full <- c(shining_list, year = 1980)

# Have a look at shining_list_full
str(shining_list_full)
```

*** =sct
```{r}
msg = "Do not remove or change the definition of `shining_list`, which is pre-loaded in the workspace!"
test_object("shining_list", undefined_msg = msg, incorrect_msg = msg)
test_object("shining_list_full", eq_condition = "equal",
            incorrect_msg = paste("Have you correctly extended `shining_list` with an element named `year`,",
                                  "whose value is 1980? You can use `c(shining_list, year = 1980)`"))
test_function("str", "object", incorrect_msg = "Don't forget to display the structure of `shining_list_full` with `str()`.")
success_msg("Great! This was the last exercise on R lists! You now have a solid basis in the R programming language, but there's so much more to learn. Check out all the other DataCamp courses and become a true data science expert!")
```


